This is an assignment

```python
def hello_world():
    return "Hello World!"
```

<!-- grade_id: the answer, points: 0, grade: false, solution: true -->
```python
def hello_world():
    return "Hello"
```

<!-- grade_id: test, points: 10, grade: true, solution: false -->
```python
assert hello_world() == "Hello!"
```

```python

```